# Exercise
Write a function which receives a positive integer number and prints the sum of all even numbers smaller than the parameter read. You can use `numpy`, but you can also use a simple loop.

<!-- grade_id: cell-48df8f2b83dee5a8, points: 0, grade: false, solution: true -->
```python
import numpy as np
def sum_smaller_even(n):
    ### BEGIN SOLUTION
    result = 0
    for ii in range(2, n, 2):
        result += ii
    return result
    # return np.sum(np.arange(2, n, 2))
    ### END SOLUTION
```

<!-- grade_id: cell-4fe77a6d354d52e1, points: 0.5, grade: true, solution: false -->
```python
assert 0 == sum_smaller_even(-10)
assert 0 == sum_smaller_even(0)
assert 0 == sum_smaller_even(2)
assert 2 == sum_smaller_even(3)
assert 2 == sum_smaller_even(4)
assert 6 == sum_smaller_even(5)
assert 6 == sum_smaller_even(6)
assert 30 == sum_smaller_even(12)
assert 42 == sum_smaller_even(13)
assert 42 == sum_smaller_even(14)
assert 2450 == sum_smaller_even(100)
### BEGIN HIDDEN TESTS
assert 2756 == sum_smaller_even(105)
assert 13340 == sum_smaller_even(231)
### END HIDDEN TESTS
```

# Exercise
Write a function (called `sum1`) that computes the mathematical sum $$\sum_{k=1}^M \frac{1}{k},$$ as a function of $M$. Write another function (calle `diff`) that receives to parameters, `m` and `n`, and computes the relative difference between `sum1(m)` and `sum2(n)`. The relative difference between two numbers $a$ and $b$ is defined as $|\frac{a-b}{a}|$ (assuming $|a| > 0$). You can use the `math abs` or `numpy abs` functions and so on. 

<!-- grade_id: cell-2be48593e458a7b6, points: 0, grade: false, solution: true -->
```python
import numpy as np
def sum1(m):
    ### BEGIN SOLUTION
    return np.sum(1.0/np.arange(1, m+1))
    ### END SOLUTION

def diff(m, n):
    ### BEGIN SOLUTION
    return np.abs(1.0 - sum1(m)/sum1(n))
    ### END SOLUTION    
```

<!-- grade_id: cell-f2cab27fa5a5cef0, points: 0.5, grade: true, solution: false -->
```python
assert np.isclose(0.0, sum1(0))
assert np.isclose(1.0, sum1(1))
assert np.isclose(7.485470860550345, sum1(1000))
assert np.isclose(10.644712286201385, sum1(23564))
assert np.isclose(1.0, diff(0, 10))
assert np.isclose(0.6688607232024465, diff(1, 11))
assert np.isclose(0.1858865473624568, diff(10, 20))
assert np.isclose(3.9856430361639994e-05, diff(23564, 23574))
### BEGIN HIDDEN TESTS
assert np.isclose(1.5, sum1(2))
assert np.isclose(2.283333333333333, sum1(5))
assert np.isclose(2.9289682539682538, sum1(10))
assert np.isclose(9.787706026045383, sum1(10001))
assert np.isclose(0.5166296601992537, diff(2, 12))
assert np.isclose(0.3118819291879539, diff(5, 15))
assert np.isclose(0.0013268618684117994, diff(1000, 1010))
assert np.isclose(0.00010209220383317064, diff(10001, 10011))
### END HIDDEN TESTS
```

# Exercise
Write a function that computes the central derivative of a given function, $$ f'(x) \simeq \frac{f(x+h) - f(x-h)}{2h}.$$ Use, as a model, the function $sin(x)$.  Write another function , called `numdiff`, than computes the relative difference between the exact and the numerical derivative 

<!-- grade_id: cell-ac4d4f37c7341bad, points: 0, grade: false, solution: true -->
```python
import numpy as np
def f(x):
    return np.sin(x)

def centraldiff(x, h):
    ### BEGIN SOLUTION
    return (f(x+h) - f(x-h))/(2*h)
    ### END SOLUTION

def numdiff(x, h):
    ### BEGIN SOLUTION
    return np.abs(1.0-centraldiff(x, h)/np.cos(x))
    ### END SOLUTION
```

<!-- grade_id: cell-91b0fd2cf74990a5, points: 0.5, grade: true, solution: false -->
```python
assert np.isclose(0.0411489227915941, numdiff(1.2478, 0.5))
assert np.isclose(0.0016658335317168982, numdiff(1.2478, 0.1))
assert np.isclose(4.166661485727197e-06, numdiff(1.2478, 0.005))
assert np.isclose(1.6666681634802671e-07, numdiff(1.2478, 0.001))
assert np.isclose(4.166687694873161e-08, numdiff(1.2478, 0.0005))
assert np.isclose(2.1800006244632186e-11, numdiff(1.2478, 1e-05))
assert np.isclose(1.2246484937250557e-08, numdiff(1.2478, 1e-08)) # NOTICE: ERRORS INCREASING
assert np.isclose(4.021997401792987e-08, numdiff(1.2478, 1e-09))
assert np.isclose(0.0031372791395899435, numdiff(1.2478, 1e-14))
assert np.isclose(0.04932917985306329, numdiff(1.2478, 1e-15))
### BEGIN HIDDEN TESTS
assert np.isclose(0.0004166145864309412, numdiff(1.2478, 0.05))
assert np.isclose(1.6666583323776685e-05, numdiff(1.2478, 0.01))
assert np.isclose(1.6674979175235194e-09, numdiff(1.2478, 0.0001))
assert np.isclose(4.135495279555812e-10, numdiff(1.2478, 5e-05))
### END HIDDEN TESTS
```

# Exercise
Write a function that receives a positive integer number and returns the sum of the prime numbers less or equal than the number. You should also define a function that checks if a given number is prime

<!-- grade_id: cell-cfaa7f286621ee36, points: 0, grade: false, solution: true -->
```python
import numpy as np
def sumprimeslessthan(n):
    ### BEGIN SOLUTION
    suma = 0
    for ii in range(1, n+1):
        if True == isprime(ii):
            suma += ii
    return suma
    ### END SOLUTION
    
def isprime(m):
    ### BEGIN SOLUTION
    if m <= 1: return False
    for ii in range(2, int(np.sqrt(m))+1):
        if m%ii == 0:
            return False
    return True
    ### END SOLUTION
```

<!-- grade_id: cell-611da4f66c6be786, points: 0.5, grade: true, solution: false -->
```python
assert False == isprime(1)
assert True == isprime(2)
assert True == isprime(3)
assert False == isprime(4)
assert True == isprime(5)
assert False == isprime(9)
assert False == isprime(12)
assert True == isprime(100267)
assert 0 == sumprimeslessthan(1)
assert 2 == sumprimeslessthan(2)
assert 5 == sumprimeslessthan(3)
assert 5 == sumprimeslessthan(4)
assert 10 == sumprimeslessthan(5)
assert 17 == sumprimeslessthan(9)
assert 28 == sumprimeslessthan(12)
assert 456299080 == sumprimeslessthan(100267)
### BEGIN HIDDEN TESTS
assert False == isprime(123)
assert False == isprime(100011)
assert True == isprime(100019)
assert 1593 == sumprimeslessthan(123)
assert 454496540 == sumprimeslessthan(100011)
assert 454596559 == sumprimeslessthan(100019)
### END HIDDEN TESTS
```